# Exercise 1: Profile Data

**Aim:** To work with vertical profile data and make some standard calculations.

**Required Outputs:**
Four figures should be generated (replace YourName with your actual last name):
- ex1fig1-YourName-Messfern.png: Showing the temperature and practical salinity
- ex1fig2-YourName-Messfern.png: Showing the conservative temperature and absolute salinity profiles
- ex1fig3-YourName-Messfern.png: Showing a basic TS profile
- ex1fig4-YourName-Messfern.png: Showing a TS profile focused on waters denser than 27.

**Fixed File Paths (DO NOT CHANGE):**
- NetCDF files: `./data/MSM121_054_1db.nc` and `./data/MSM121_054_1db_edited.nc`
- Figures directory: `./figures/`

---

<!-- grade_id: student_info, points: 0, grade: false, solution: true -->
**Complete your information here:**

- **Your Name:** [REPLACE WITH YOUR ACTUAL NAME]
- **Date:** [REPLACE WITH TODAY'S DATE]
- **Student ID:** [REPLACE WITH YOUR STUDENT ID]

**Replace the placeholder text above with your actual information.**

## Importing Required Packages

```python
# Importing required packages here
import matplotlib.pyplot as plt
import gsw 
import numpy as np
import xarray as xr
import os

from matplotlib.ticker import ScalarFormatter

# Import seasenselib for CTD data processing
from seasenselib.readers import SbeCnvReader, NetCdfReader
from seasenselib.writers import NetCdfWriter
```

## File Paths Setup

<!-- grade_id: file_setup, points: 0, grade: false, solution: true -->
```python
# Input file paths hLLO
cnv_file_with_path = '../data/MSM121_054_1db.cnv'

# Output file paths for netCDF (DO NOT CHANGE THESE)
netcdf_file_with_path = '../data/MSM121_054_1db.nc'
netcdf_file_with_path_edited = '../data/MSM121_054_1db_edited.nc'

# YOUR CODE HERE
# Set up the figures directory - DO NOT CHANGE THIS PATH
# The autograder expects figures to be saved in '../figures/'
figdir = '../figures/'
raise NotImplementedError()

# Create figures directory if it doesn't exist
if not os.path.exists(figdir):
    os.makedirs(figdir)
    
print(f"Figures will be saved to: {figdir}")
print(f"Data file: {cnv_file_with_path}")
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[2], line 12
      8 # YOUR CODE HERE
      9 # Set up the figures directory - DO NOT CHANGE THIS PATH
     10 # The autograder expects figures to be saved in '../figures/'
     11 figdir = '../figures/'
---> 12 raise NotImplementedError()
     14 # Create figures directory if it doesn't exist
     15 if not os.path.exists(figdir):

NotImplementedError:
```

<!-- grade_id: test_file_setup, points: 1, grade: true, solution: false -->
```python
# Test that setup is complete
assert 'figdir' in locals(), "figdir variable not defined"
assert figdir.endswith('/'), "figdir should end with '/'"
# Note: CNV file will be downloaded automatically when needed, so we don't test for existence here
print("✓ File setup test passed!")
```

## Transforming to xarray

### Convert CNV file to netCDF

<!-- grade_id: cnv_conversion, points: 0, grade: false, solution: true -->
```python
# Using seasenselib for CTD data processing
# Data from Seabird sensors typically come in as .cnv files.  
# These files can be converted to netCDF format using the seasenselib package.

# YOUR CODE HERE
# Using seasenselib:   
# Read CTD data from CNV file using SbeCnvReader
# Get the dataset using reader.get_data()
# Name the dataset ctd_data
raise NotImplementedError()

```

### Write to netCDF


```python
# If there are problems with saving because of attributes, then uncomment this line
#ctd_data = ctd_data.drop_attrs()
#
# Write dataset with CTD data to netCDF file using NetCdfWriter
# Only write if the file doesn't already exist
raise NotImplementedError()

print(f"NetCDF conversion complete: {netcdf_file_with_path}")
```

<!-- grade_id: test_cnv_conversion, points: 2, grade: true, solution: false -->
```python
# Test that conversion worked
assert os.path.exists(netcdf_file_with_path), "NetCDF file was not created"
print("✓ CNV to NetCDF conversion test passed!")
```

### Open file and show contents

<!-- grade_id: data_loading, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# load data using xr.open_dataset()
raise NotImplementedError()

# print data
print(ctd_ds.info)
```

<!-- grade_id: test_data_loading, points: 1, grade: true, solution: false -->
```python
# Test that data loaded correctly
assert 'ctd_ds' in locals(), "ctd_ds variable not found"
assert hasattr(ctd_ds, 'data_vars'), "ctd_ds should be an xarray Dataset"
print("✓ Data loading test passed!")
```

## Understanding the data format

The data stored within the netcdf file include measured parameters as `variables`, for example:

- `pressure` is pressure in decibars
- `temperature` is temperature in degrees Celcius
- `salinity` is salinity from the secondary sensor
- `latitude` is the position where the data were take from in degrees North (positive) or south (negative)
- `longitude` is the longitude with degrees east (positive) or west (negative)

The list of `Attributes` include details like the single value for latitude/longitude, or the date the file was created.

```python
# Here we are defining the names of the variables in the netCDF file
# If your netcdf file calls salinity "PSAL", you would change the sal_string to 'PSAL'
sal_string = 'salinity'
temp_string = 'temperature'
pres_string = 'pressure'
lon_string = 'longitude'
lat_string = 'latitude'
SA_string = 'absolute_salinity'
CT_string = 'conservative_temperature'
# You could also define these pairs in a dictionary (https://www.w3schools.com/python/python_dictionaries.asp)
```

# Plot the temperature and salinity

To plot in Python, we often use the package `matplotlib` which you've imported above. This includes a variety of functions which can be used to generate figures in python.

<!-- grade_id: figure1, points: 0, grade: false, solution: true -->
```python
# To initialise a figure, use matplotlib.plt.figure() which takes inputs within the brackets
# Since matplotlib is long to write, in the import cell above, we've imported it as plt
plt.figure(1, figsize=(12, 8)) 

# In case the figure 1 already exists, clear it using clf()
plt.clf()

# Within the figure window, we'll actually have two separate sets of axes.  These are called 'subplots' in matplotlib
# We'll have two subplots, one next to the other, so we'll have a 1x2 grid of subplots
plt.subplot(1,2,1)

# To plot salinity, we need to reference the salinity variable in the dataset.  This is done using the syntax ctd_ds['PSAL2']
# Here, ctd_ds is the name of the dataset, and 'PSAL2' is the name of the variable.  The variable name is case sensitive.
plt.plot(ctd_ds[sal_string], ctd_ds[pres_string], color='blue')

# Now we'd like to annotate our figure.  We can do this using the plt.xlabel() and plt.ylabel() functions
# These functions take a string as an argument, which is the text that will be displayed on the x and y axes
# An optional second argument can be used to specify the fontsize.  IMPORTANT: choose values to make the text readable.
plt.xlabel('Practical salinity' , fontsize=12)
plt.ylabel('Pressure [dbar]', fontsize=12)

# Since pressure decreases with depth, we want to invert the y-axis so that the surface is at the top of the plot
plt.gca().invert_yaxis()

# Here we specify where exactly the x-ticks should be placed.  We want them at 34, 34.5, and 35
x_ticks = [34, 34.5, 35]
plt.xticks(x_ticks, fontsize=12)

# We don't specify for the y-axis, but we do want the text size to be readable.
plt.yticks(fontsize=12)

# Finally, we can add a title to the plot using plt.title().  This function also takes a string as an argument.
plt.title('Salinity profile', fontsize=12)
plt.grid(True)

# YOUR TURN: Now, let's plot the temperature profile on the right-hand side of the figure
plt.subplot(1,2,2)

# YOUR CODE HERE
# Complete the temperature plot following the same pattern as the salinity plot
# Plot temperature vs pressure, add labels, title, formatting
# Remember to invert the y-axis and add grid
raise NotImplementedError()

# This code saves the figure to a file.  The file type is determined by the extension of the filename.
# IMPORTANT: Change YourName to your actual last name
plt.savefig(figdir + 'ex1fig1-YourName-Messfern.png')
plt.show()
```

<!-- grade_id: test_figure1, points: 3, grade: true, solution: false -->
```python
# Test that Figure 1 was created correctly
import glob
fig1_files = glob.glob(figdir + 'ex1fig1-*-Messfern.png')
assert len(fig1_files) > 0, "Figure 1 (ex1fig1-YourName-Messfern.png) was not created"
for fig_file in fig1_files:
    assert 'YourName' not in fig_file, f"Replace 'YourName' with your actual name in: {fig_file}"
print(f"✓ Figure 1 test passed: {fig1_files[0]}")
```

# Checklist for figure details to change:

- [ ] Change the shallow limit to be exactly zero
- [ ] Change the gridlines to be dashed (or dotted)
- [ ] Add additional ticks on the x-axis to show more detail
- [ ] Reduce the ticks and gridlines on the y-axis to be every 1000m
- [ ] Change the title to also include the latitude/longitude of the profile
- [ ] Add empty square brackets after Practical Salinity [ ] to indicate that there are no units
- [ ] Make the figure size about 1/3rd the height of A4 paper, and the text width to fit on A4 paper with margins.
- [ ] Change the font to Times New Roman or another serif font.

## Calculating TEOS-10 parameters

While data are collected as practical salinity and in situ temperature, in Messmethoden, we'll learn that we plot or represent the data by the quantities "Absolute salinity" and "conservative temperature". These are better tracers of the properties within the ocean.

To convert from practical to absolute salinity and temperature to conservative, we use the TEOS-10 toolbox, sometimes called the "Gibbs Seawater Toolbox". In Python, you've loaded this as `import gsw`.

https://teos-10.github.io/GSW-Python/conversions.html

<!-- grade_id: teos10_calculations, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# Calculate absolute salinity
# SA_2 = gsw.conversions.SA_from_SP(...)

# Calculate conservative temperature
# CT_2 = gsw.conversions.CT_from_t(...)

# Add the new variables to the dataset
# Note: Use .values to extract the data from the DataArray
# ctd_ds[SA_string] = (ctd_ds[sal_string].dims, SA_2.values)
# ctd_ds[SA_string].attrs['long_name'] = 'Absolute Salinity'
# ctd_ds[SA_string].attrs['units'] = 'g/kg'

# Follow the same pattern for conservative temperature
raise NotImplementedError()

# Check the data.  Note that the new variables are now in the dataset.
print(ctd_ds.info)
```

<!-- grade_id: test_teos10, points: 3, grade: true, solution: false -->
```python
# Test TEOS-10 calculations
assert SA_string in ctd_ds.data_vars, f"Missing {SA_string} variable in dataset"
assert CT_string in ctd_ds.data_vars, f"Missing {CT_string} variable in dataset"
print("✓ TEOS-10 calculations test passed!")
```

You can also calculate and add the field to the data structure in a single line

```python
# Calculate depth (optional)
#ctd_ds['DEPTH'] = gsw.z_from_p(ctd_ds['PRES'], ctd_ds['LATITUDE']) # convert pressure to depth
```

## Plot absolute and conservative temperature profiles

<!-- grade_id: figure2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# Plot the absolute salinity and conservative temperature profiles against depth in meters.
# Create a similar two-panel figure to Figure 1, but using the TEOS-10 variables
# Left panel: Absolute salinity vs pressure
# Right panel: Conservative temperature vs pressure
raise NotImplementedError()

# IMPORTANT: Change YourName to your actual last name
plt.savefig(figdir + 'ex1fig2-YourName-Messfern.png')
plt.show()
```

<!-- grade_id: test_figure2, points: 3, grade: true, solution: false -->
```python
# Test that Figure 2 was created correctly
fig2_files = glob.glob(figdir + 'ex1fig2-*-Messfern.png')
assert len(fig2_files) > 0, "Figure 2 (ex1fig2-YourName-Messfern.png) was not created"
for fig_file in fig2_files:
    assert 'YourName' not in fig_file, f"Replace 'YourName' with your actual name in: {fig_file}"
print(f"✓ Figure 2 test passed: {fig2_files[0]}")
```

### Saving xarray to netCDF file

Since we've manipulated the data, and may want to use the new variables "Absolute salinity" and "Conservative temperature" again later, let's save the new data file.

```python
ctd_ds.to_netcdf(netcdf_file_with_path_edited)
```

## T-S Diagram

A T-S diagram has salinity on the x-axis and temperature on the y-axis, and contours of density (sigma_0) added.

<!-- grade_id: figure3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# Follow the original template pattern to create T-S diagram
# Reference: https://stackoverflow.com/questions/41634179/matplotlib-contour-lines-are-not-closing-up

# On a TS diagram, the X axis is salinity and the Y axis is temperature.  
# As a helpful reminder, we can create variables X and Y.
# x = ctd_ds[SA_string]
# y = ctd_ds[CT_string]

# A typical TS diagram also has contours of potential density using gsw package.
# Use np.meshgrid and gsw.density.sigma0() to create density contours
# Plot the data points and add contour lines with labels
raise NotImplementedError()

# IMPORTANT: Change YourName to your actual last name
plt.savefig(figdir + 'ex1fig3-YourName-Messfern.png')
plt.show()
```

<!-- grade_id: test_figure3, points: 3, grade: true, solution: false -->
```python
# Test that Figure 3 was created correctly
fig3_files = glob.glob(figdir + 'ex1fig3-*-Messfern.png')
assert len(fig3_files) > 0, "Figure 3 (ex1fig3-YourName-Messfern.png) was not created"
for fig_file in fig3_files:
    assert 'YourName' not in fig_file, f"Replace 'YourName' with your actual name in: {fig_file}"
print(f"✓ Figure 3 test passed: {fig3_files[0]}")
```

## Enhanced T-S Diagram (Figure 4)

Repeat the plot here, but with the following enhancements:
- Plot the contours first, then the data points (makes data points more visible)
- Try removing the grid lines
- Change the aspect ratio of the plot to make it square
- Change the color of the line to red
- Change the axes limits to focus on the data denser than 27.00 kg/m³

<!-- grade_id: figure4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# Repeat the T-S plot from Figure 3, but with these modifications:
# - Plot the contours first, then the data points (makes data points more visible)
# - Remove grid lines
# - Change the aspect ratio of the plot to make it square
# - Change the color of the line to red
# - Change the axes limits to focus on data denser than 27.00 kg/m^3
raise NotImplementedError()

# Save the figure as a .png file in the figures directory
# IMPORTANT: Change YourName to your actual last name
plt.savefig(figdir + 'ex1fig4-YourName-Messfern.png')
plt.show()
```

<!-- grade_id: test_figure4, points: 3, grade: true, solution: false -->
```python
# Test that Figure 4 was created correctly
fig4_files = glob.glob(figdir + 'ex1fig4-*-Messfern.png')
assert len(fig4_files) > 0, "Figure 4 (ex1fig4-YourName-Messfern.png) was not created"
for fig_file in fig4_files:
    assert 'YourName' not in fig_file, f"Replace 'YourName' with your actual name in: {fig_file}"
print(f"✓ Figure 4 test passed: {fig4_files[0]}")
```

<!-- grade_id: analysis_questions, points: 0, grade: false, solution: true -->
## Analysis Questions

**Note:** These questions are for your reflection and will be discussed in class after the assignment deadline. You do not need to submit written answers - they are not part of the automated grading.

**Question 1:** Compare the practical salinity vs. absolute salinity profiles (Figures 1 and 2). What differences do you observe and why are these important for oceanographic analysis?

**Question 2:** Examine your T-S diagrams (Figures 3 and 4). What do the density contours tell you about the water masses in this profile? Which parts of the water column are most/least dense?

**Question 3:** Based on the location of this CTD cast (North Atlantic, 49°42.5'N, 45°00.0'W) and the T-S characteristics you observe, what can you infer about the water masses present? Consider both surface and deep water properties.

**Optional Notes for Class Discussion:**
Feel free to jot down thoughts here, but this section is not graded:

1. [Your observations for Question 1]

2. [Your observations for Question 2]

3. [Your observations for Question 3]

## Conclusion

Congratulations! You have successfully completed this comprehensive oceanographic data analysis exercise. You have:

- Converted Sea-Bird CNV data to netCDF format using seasenselib
- Calculated TEOS-10 oceanographic parameters (Absolute Salinity, Conservative Temperature)
- Created professional oceanographic profile plots (Figures 1 & 2)
- Generated Temperature-Salinity diagrams with density contours (Figures 3 & 4)
- Analyzed real research data from the North Atlantic Ocean

This exercise demonstrates the complete workflow used by oceanographers for analyzing CTD data, from raw sensor data to scientific interpretation.

## Grading

**This assignment uses a tiered PASS system:**

✅ **PASS (Basic Requirements):**
- All required files created (NetCDF files, 4 figures)
- All automated tests pass (data valid, figures contain content)
- Personal information completed
- Code executes without errors

🌟 **PASS PLUS (Professional Quality):**
All PASS requirements PLUS:
- Sensible legends and titles on all figures
- Good choices for figure plotting (appropriate aspect ratios, professional color scales, clear line styles)
- All axes have proper labels with units (e.g., "Temperature (°C)", "Pressure (dbar)", "Salinity (g/kg)")
- Thoughtful formatting (readable fonts, appropriate figure sizes)
- Professional T-S diagram styling with proper density contour labels

❌ **FAIL Conditions:**
- Missing required outputs
- Test failures (invalid data, empty figures)
- Code execution errors
- Template placeholders not replaced

**Analysis Questions:** The analysis questions above are for reflection and will be discussed in class after the assignment deadline.